# Homework 1: Data Types, Functions and Conditionals (29 points)

name: Zhekai Liu

email: zhekail@umich.edu

This homework assignment took me 10 hours in total to complete. (Please help us to gauge the difficulty of the assignment.)

What is your primary reason for taking this class? (List just THE most important reason).

I majored in math as an undergraduate. In master's study, I would like to learn something about computer science and programming to improve my skills.

## Collaboration Disclosure

In the cell below, please list *everyone* with whom you discussed any of the homework problems, excluding only the GSIs and the course instructor. 

If you did not discuss the homework with anyone else, write __"I did not discuss this homework with anyone."__

Even if you discuss questions with other, the code you submit must be only yours. All work is checked with the [MOSS plagiarism detector](https://theory.stanford.edu/~aiken/moss/).)

I did not discuss this homework with anyone.

## Submission Instructions
Your homework solutions should be written entirely in this Jupyter notebook file. Once it contains your solutions, you should submit this notebook through Canvas. 


Before submitting, please make sure to __Cells->Run All__ executes without errors; errors in your code translate directly to point deductions. 
In general, you don't need to do explicitly raise errors (e.g. with the ```raise``` function) if we don't ask you to in the problem statement.
However, even in cases where we ask you to check for errors, your submission should not contain any examples of your functions actually raising those errors.

Note that many parts of this homework where you are expected to type in code will have ```NotImplementedError()``` as a placeholder. You need to delete this function and replace it with your own code.

## Importing libaries

Unless specifically prohibited, you can import any Python library that comes included with Python (e.g. ```math```, ```datetime```) or that has been used in the lecture or homework prompts. You do not need any other libraries to complete the homework. Note that the graders will not install any additional libraries to grade an assignment.


<!-- ## Coding style
Please follow the [PEP8 guidelines](https://www.python.org/dev/peps/pep-0008/). The `pycodestyle_magic` package will remind you if you are not. To install it, run
```
pip install flake8 pycodestyle_magic
```
Once it is installed the following two Jupyter directives will enable syntax checking: -->

## Homework tips 

1. **Start early!** If you run into trouble installing things or importing packages, it’s
best to find those problems well in advance, not the night before your assignment is
due when we cannot help you!

2. **Make sure you back up your work!** At a minimum, do your work in a Dropbox
folder. Better yet, use git, which is well worth your time and effort to learn.

3. **Be careful to follow directions!** Remember that Python is case sensitive. If
you are ask you to define a function called my_function and you define a function
called My_Function, you will not receive full credit. You may want to copy-paste
the function names below to make sure that the functions in your notebook match.


## Nbgrader

We will be using `nbgrader` to grade your jupyter notebook. You will notice some `read-only` cells in the assignment that contain `assert` statements. These are tests that your code must pass for your solution to be correct. If any of the tests fail, you will get an python error and not get points for that question. 

**Note:** The tests shown not are not comprehensive; additional tests will be used at grading time. You are encouraged to read the problem carefully and verify your code covers all possible cases.

**Be careful:** If a jupyter notebook cell takes longer than `60s` to run, the autograder will not grade it and you will receive zero credit for that question.

## Problem 1: Severance textbook exercises (8 points) 
Complete the following exercises from the textbook which you can find [here](https://www.py4e.com/html3/) for free. 

Note that some questions will ask you to get input from the user via the ```input```
function. Since Jupyter has no command-line interface, you will instead implement the program as a function. The answer cells you find below will provide a template to help you get started.

- Ch 1, Exercises 3,5,6,8  
- Ch 2, Exercises 3,4  
- Ch 3, Exercises 1, 2  
- Ch 4, Exercise 6  
- Ch 6, Exercise 1  

---
### Chapter 1, Exercises (2 points) 

#### Exercise 3 
What is the difference between a compiler and an interpreter?

<!-- grade_id: ch1-ex3-ans, points: 0, grade: false, solution: true -->
```python
# Enter your answer as a python string in a variable called 'ans'. 
# For example: "ans = 'This is a sample answer.'"
# You can use markdown formatting. 
# YOUR CODE HERE
ans = "The compiler needs to take the source code file of the entire program and convert it into machine language. An interpreter reads a program's source code line by line, parses the source code, and executes the instructions instantly. Instead of translating the entire program into machine code before execution, it reads and executes as it goes."
```

<!-- grade_id: ch1-ex3-test, points: 0.5, grade: true, solution: false -->
```python
# Make sure you answer is rendered as markdown below. 
from IPython.display import Markdown as md
assert type(ans) == str
md(ans)
```

Output:
```
<IPython.core.display.Markdown object>
```

#### Exercise 5
What is wrong with the following code:
```
>>> primt 'Hello world!'
File "<stdin>", line 1
  primt 'Hello world!'
                     ^
SyntaxError: invalid syntax
>>>
```

<!-- grade_id: ch1-ex5-ans, points: 0, grade: false, solution: true -->
```python
# Enter your answer as a python string in a variable called 'ans'. 
# You can use markdown formatting. 
# YOUR CODE HERE
ans = "‘print’ was misspelled as ‘primt’."
```

<!-- grade_id: ch1-ex5-test, points: 0.5, grade: true, solution: false -->
```python
# Make sure you answer is rendered as markdown below. 
from IPython.display import Markdown as md
assert type(ans) == str
md(ans)
```

Output:
```
<IPython.core.display.Markdown object>
```

#### Exercise 6
Where in the computer is a variable such as `x` stored after the following Python line finishes?

a) Central processing unit

b) Main Memory

c) Secondary Memory

d) Input Devices

e) Output Devices

<!-- grade_id: cell-b0f556c01b7e2276, points: 0, grade: false, solution: true -->
```python
# Enter your answer as a python string in a variable 'ans'. 
# Should be one of ['a', 'b', 'c', 'd', 'e']
# YOUR CODE HERE
ans = 'b'
```

<!-- grade_id: cell-3176d2d916422842, points: 0.5, grade: true, solution: false -->
```python
assert type(ans) == str
assert ans in ['a','b','c','d','e']
```

#### Exercise 8
Explain each of the following using an example of a human capability: (1) Central processing unit, (2) Main Memory, (3) Secondary Memory, (4) Input Device, and (5) Output Device. For example, “What is the human equivalent to a Central Processing Unit”?

<!-- grade_id: ch1-ex8-ans, points: 0, grade: false, solution: true -->
```python
# Enter your answer as a python string in a variable called 'ans'. 
# You can use markdown formatting. 
# YOUR CODE HERE
ans = "(1) Brain, (2) Short-term memory, (3) Long-term memory, (4) Sensory (e.g. visual, auditory, tactile), (5) Communication/action (e.g. speaking, writing, moving)"
```

<!-- grade_id: ch1-ex8-test, points: 0.5, grade: true, solution: false -->
```python
# Make sure you answer is rendered as markdown below. 
from IPython.display import Markdown as md
assert type(ans) == str
md(ans)
```

Output:
```
<IPython.core.display.Markdown object>
```

### Chapter 2, Exercises (2 points) 

#### Exercise 3
 Write a function which computes gross pay given the **hours** worked and the **rate_per_hour**. 
 
 (Note: You do not need to use the input() function as in the textbook question.)

<!-- grade_id: cell-312caeaaf6abe1a9, points: 0, grade: false, solution: true -->
```python
def get_gross_pay(hours, rate_per_hour):
    # YOUR CODE HERE
    return hours*rate_per_hour
```

<!-- grade_id: cell-458cb3b5b6333645, points: 1, grade: true, solution: false -->
```python
assert get_gross_pay(3, 10) == 30
```

#### Exercise 4
Assume that we execute the following assignment statements:
```
width = 17
height = 12.0
```
For each of the following expressions, write the value of the expression and the type (of the value of the expression).

1. `width//2`

2. `width/2.0`

3. `height/3`

4. `1 + 2 * 5`

<!-- grade_id: cell-e69aa106b6a89086, points: 0, grade: false, solution: true -->
```python
# Enter you answer as python list of tuples (value, type) in a variable 'ans'.
# For example, "ans = [(1,int), (2.0, float)]" (not actual answer)
# where the order of the list corresponds to order above.

# YOUR CODE HERE
ans = [(8,int), (8.5, float), (4.0, float), (11,int)]
```

<!-- grade_id: cell-796c8c2828e551ce, points: 1, grade: true, solution: false -->
```python
assert ans[0] == (8,int)
```

---
### Chapter 3, Exercises (2 points) 

#### Exercise 1 
Rewrite your `get_gross_pay` function from ch2, exercise 3 to give the employee 1.5 times the hourly rate for hours worked above 40 hours.

<!-- grade_id: cell-8f610a982f866fac, points: 0, grade: false, solution: true -->
```python
def get_gross_pay(hours, rate_per_hour):
    # YOUR CODE HERE
    extra_time = hours-40
    extra_wage = extra_time*rate_per_hour*1.5
    if extra_wage > 0:
        return extra_wage + 40*rate_per_hour
    else:
        return hours*rate_per_hour
```

<!-- grade_id: cell-537d9d1c5dcd290f, points: 1, grade: true, solution: false -->
```python
assert get_gross_pay(0, 1000)==0
assert get_gross_pay(2, 10)==20
```

#### Exercise 2
Rewrite your `get_gross_pay` program from ch2, exercise 3 using `try` and `except` so that your program handles non-numeric input gracefully by printing an informative message and returning nothing.

```
>>> get_gross_pay(20, 'nine')
Error, please enter numeric input

>>> get_gross_pay('forty', 'nine')
Error, please enter numeric input
```

<!-- grade_id: cell-b57ce5b94cb6cf92, points: 0, grade: false, solution: true -->
```python
def get_gross_pay(hours, rate_per_hour):
    # YOUR CODE HERE
    try:
        hours = float(hours)
        rate_per_hour = float(rate_per_hour)
    except:
        print("Error, please enter numeric input")
        return
    
    extra_time = hours-40
    extra_wage = extra_time*rate_per_hour*1.5
    if extra_wage > 0:
        return extra_wage + 40*rate_per_hour
    else:
        return hours*rate_per_hour
```

<!-- grade_id: cell-582c5cae10097fd1, points: 1, grade: true, solution: false -->
```python
try: 
    get_gross_pay('nine', 'five')
    get_gross_pay('nine', 5)
    get_gross_pay(5, 5)
    get_gross_pay(5, 'nine')
    assert True
except: 
    assert False
```

Output:
```
Error, please enter numeric input
Error, please enter numeric input
Error, please enter numeric input
```

### Chapter 4, Exercises (1 point)

#### Exercise 1 
What is the purpose of the `def` keyword in Python?

a) It is slang that means “the following code is really cool”

b) It indicates the start of a function

c) It indicates that the following indented section of code is to be stored for later

d) b and c are both true

e) None of the above

<!-- grade_id: ch4-ex1-ans, points: 0, grade: false, solution: true -->
```python
# Enter your answer as a python string in a variable 'ans'. 
# Should be one of ['a', 'b', 'c', 'd', 'e']
# YOUR CODE HERE
ans = 'd'
```

<!-- grade_id: ch4-ex1-test, points: 1, grade: true, solution: false -->
```python
assert type(ans) == str
assert ans in ['a','b','c','d','e']
```

### Chapter 6, Exercises (1 point) 

#### Exercise 6
Read the documentation of the string methods at https://docs.python.org/library/stdtypes.html#string-methods You might want to experiment with some of them to make sure you understand how they work. `strip` and `replace` are particularly useful.

<!-- grade_id: cell-9a484821ec48c9fc, points: 0, grade: false, solution: true -->
```python
# Create a variable 'ans' with the string: 'I read the documentation.'
# YOUR CODE HERE
ans = 'I read the documentation.'
```

<!-- grade_id: cell-4dd3c9dad6369e20, points: 1, grade: true, solution: false -->
```python
assert ans == 'I read the documentation.'
```

## Problem 2: Euclid's Algorithm (3 points)

Euclid’s algorithm (https://en.wikipedia.org/wiki/Euclidean_algorithm) is a method
for finding the greatest common divisor (GCD) of two integers. Recall that the GCD of
two integers $m$ and $n$ is the largest number that divides both $m$ and $n$.

---
### Part 1 (2 points)

The __Wikipedia__ page above includes several pseudocode implementations of Euclid’s
algorithm. Choose one of these, and use it to implement a function `gcd`, which takes
two integers as its arguments and returns their GCD. You may assume that both
inputs are postive integers (i.e. $m,n \in \mathbb Z$ and $m, n > 0$), so there is no need to include any error checking in your function.

**Note:** this is one of the rare occasions where you have my explicit permission to
look up your answer. Unless otherwise stated (e.g., as in this problem), looking up
solutions on Wikipedia or in any other non-class resource will be considered cheating!

<!-- grade_id: gcd, points: 0, grade: false, solution: true -->
```python
def gcd(a,b):
    """Compute the great common divisor (GCD) of two integers a,b. 
    """
    # YOUR CODE HERE
    while b != 0:
        t = b
        b = a % b
        a = t
    return a
```

<!-- grade_id: gcd_tests, points: 2, grade: true, solution: false -->
```python
"""Check that your gcd function returns the correct output for several inputs"""
assert gcd(4,5) == 1
assert gcd(20,10) == 10
```

### Part 2 (1 point) 
What does your function do if one or both of the arguments are negative?

<!-- grade_id: cell-9b7a530cbce57b8d, points: 0, grade: false, solution: true -->
```python
# Enter your answer as a python string in a variable called 'ans'. 
# You can use markdown formatting. 
# YOUR CODE HERE
ans = 'Based on my funciton, if b is positive, the result will be positive whether a is positive or not. But if b is negative, the result will always be negative. Since I use t=b and a mod b will always negative if b is negative. The absolute value of the answer will be same as the result of using absolute value of a and b.'
```

<!-- grade_id: cell-2bdb84d0d76cdb07, points: 1, grade: true, solution: false -->
```python
# Make sure you answer is rendered as markdown below. 
from IPython.display import Markdown as md
assert type(ans) == str
md(ans)
```

Output:
```
<IPython.core.display.Markdown object>
```

## Problem 3: Approximating Euler's number $e$ (8 points) 

The base of the natural logarithm, $e$, is typically deﬁned as the inﬁnite sum below:
$$
\begin{align}
    e = \sum_{k=0}^{\infty} \frac{1}{k!} = 1 + 1 + \frac{1}{2} + \frac{1}{6} + \frac{1}{24} + \dots && (1) \\
\end{align}
$$

The term $k!$ denotes the factorial of $k$, which is
$$
k! = k \cdot (k-1) \cdot (k-2) ... 3 \cdot 2 \cdot 1 
$$

By convention, $0! = 1$ For more on Euler’s number, see [here](https://en.wikipedia.org/wiki/E_(mathematical_constant)). In this problem, we will explore diﬀerent approaches to approximating this number.

### Part 1 (2 points) 

An early characterization of Euler’s number, due to Jacob Bernoulli, was as the limit of

$$ 
(1 + \frac{1}{x})^{x}
$$

as $x \rightarrow \infty$. Deﬁne a function called `euler_limit` that takes as an argument an integer $n$, and returns a `ﬂoat` that approximates $e$ by taking $x = n$ in Equation (2). You may assume that the input to your function will be a positive integer (of course, your function will still run just ﬁne if `n` is, say, a positive ﬂoat, but we will only use integer values in what follows).

<!-- grade_id: euler_limit, points: 0, grade: false, solution: true -->
```python
def euler_limit(n): 
    """Takes in a positive integer n and returns a float approximating the natural number e."""
    # YOUR CODE HERE
    e = (1+1/n)**n
    return e
```

<!-- grade_id: test_euler_limit1, points: 1, grade: true, solution: false -->
```python
"""Check that your euler_limit function returns the correct output for several inputs"""
assert euler_limit(1) == 2.0 
assert euler_limit(2) == 2.25
```

<!-- grade_id: test_euler_limit2, points: 1, grade: true, solution: false -->
```python
assert round(euler_limit(3),3) == 2.370
```

### Part 2 (1 point) 
What happens when you call `euler_limit(n)` for really huge values of `n` (say, $10^{16}$ or $10^{18}$ )? Why does this happen? 

**Hint:** the answer has to do with ﬂoating point arithmetic. Think about how $1/n$ is represented on your computer when $n$ is big.

<!-- grade_id: prob3-part2-ans, points: 0, grade: false, solution: true -->
```python
# Enter your answer as a python string in a variable called 'ans'. 
# You can use markdown formatting. 
# YOUR CODE HERE
ans = "When n is very large, 1/n becomes very small. Because of the way floating point numbers are represented in computers, when the precision of a number exceeds what a computer can represent, 1/n may become 0 or close to 0. Then (1 + 1/n) may actually be computed as 1. The result (1 + 1/n)^n may be close to 1^n = 1."
```

<!-- grade_id: prob3-part2-tests, points: 1, grade: true, solution: false -->
```python
# Make sure you answer is rendered as markdown below. 
from IPython.display import Markdown as md
assert type(ans) == str
md(ans)
```

Output:
```
<IPython.core.display.Markdown object>
```

### Part 3 (2 points) 
Deﬁne a function called `euler_sum` that takes a single non-negative integer argument `n`, and returns an approximation to $e$ based on the ﬁrst `n` terms of the sum in Equation (1). Your function should return a ﬂoat. 
You may assume that the input will be a non-negative integer, so you do not need to include error checking in your function. 

As an example, `euler_sum(4)` should return the sum of the ﬁrst four terms in Equation 1, so that `euler_sum(4)` returns $1 + 1 + 1/2 + 1/6 \approx 2.667$. 

**Note:** the sum in Equation 1 starts counting with $k = 0$ (i.e., it is “0-indexed”), while our function starts counting with $n = 1$ (i.e., it is “1-indexed”). `euler_sum(1)` should use one term from Equation (1), so that `euler_sum(1)` returns `1`. Similarly, `euler_sum(0)` should return `0`, since by convention an empty sum is equal to zero. 

**Note:** you may use the math.factorial function to compute $k!$, but I recommend, for the sake of practice, implementing the factorial function yourself, as it is a nice example of a problem that is easily solved with recursion.

<!-- grade_id: euler-sum, points: 0, grade: false, solution: true -->
```python
def euler_sum(n):
    # YOUR CODE HERE
    if n==0:
        return 0
    else:
        euler = 1
        fact = 1
        for i in range (1,n):
            fact *= i
            euler += 1/fact
        return euler
```

<!-- grade_id: test-euler-sum1, points: 1, grade: true, solution: false -->
```python
assert euler_sum(0) == 0 
assert round(euler_sum(4),3) == 2.667
```

<!-- grade_id: test-euler-sum2, points: 1, grade: true, solution: false -->
```python
assert euler_sum(3) == 2.5
```

### Part 4 (1 point) 
Deﬁne a function called `euler_approx` that takes a single argument, a positive ﬂoat `epsilon`, and uses the sum in Equation (1) to obtain an approximation of $e$ that is within `epsilon` of the true value of $e$. 

**Hint:** use a while-loop. 

**Note:** you can use the Python math module to get the true value of $e$ (up to ﬂoating point accuracy): `math.exp(1)`.

<!-- grade_id: euler-approx, points: 0, grade: false, solution: true -->
```python
import math

def euler_approx(epsilon):
    # YOUR CODE HERE
    true_euler = math.exp(1)
    n = 1
    euler = 1
    fact = 1
    error = abs(true_euler-euler)
    while error > epsilon:
        fact *= n
        euler += 1/fact
        n += 1
        error = abs(true_euler-euler)
    return euler
```

<!-- grade_id: euler-approx-tests, points: 1, grade: true, solution: false -->
```python
assert abs(euler_approx(1e0) - math.exp(1)) < 1e0
```

### Part 5 (1 point) 
For a value of interest $x$ and an approximation $\hat x$, we say the difference $\hat x - x$ is the **error** of $f(k)$.

Deﬁne functions called `print_euler_sum_error` and `print_euler_lim_error` that take a single positive integer `n` as input. For $k=1,\dots,n$, the function will print the error of the approximation of $e$ given by `euler_sum` or `euler_limit`, respectively, evaluated at $k$. Each error should be printed on its own line (for $n$ lines in total).

<!-- grade_id: prob3-ex5-func, points: 0, grade: false, solution: true -->
```python
def print_euler_sum_error(n):
    # YOUR CODE HERE
    for i in range (1, n+1):
        error = euler_sum(i) - math.exp(1)
        print(error)
    return
    
def print_euler_lim_error(n):
    # YOUR CODE HERE
    for i in range (1,n+1):
        error = euler_limit(i)-math.exp(1)
        print(error)
    return
```

<!-- grade_id: cell-fccdb9f21de162ca, points: 1, grade: true, solution: false -->
```python
# Run this cell to check your functions. 
print_euler_sum_error(10)
print()
print_euler_lim_error(10)
```

Output:
```
-1.718281828459045
-0.7182818284590451
-0.2182818284590451
-0.05161516179237857
-0.009948495125712054
-0.0016151617923787498
-0.0002262729034896438
-2.7860205076724043e-05
-3.0586177750535626e-06
-3.0288585284310443e-07

-0.7182818284590451
-0.4682818284590451
-0.3479114580886753
-0.2768755784590451
-0.22996182845904567
-0.19665545671693163
-0.17178213141833298
-0.1524973145086972
-0.13710703674584668
-0.12453936835904278
```

### Part 6 (1 point) 
Which approximation is better? 

<!-- grade_id: prob3-part6-ans, points: 0, grade: false, solution: true -->
```python
# Enter your answer as a python string in a variable called 'ans'. 
# You can use markdown formatting. 
# YOUR CODE HERE
ans = 'Equation (1) of Euler approach is better. From the error functions we could find that the error of euler sum decreased much faster than Jacob Bernoulli limits.'
```

<!-- grade_id: prob3-part6-test, points: 1, grade: true, solution: false -->
```python
# Make sure you answer is rendered as markdown below. 
from IPython.display import Markdown as md
assert type(ans) == str
md(ans)
```

Output:
```
<IPython.core.display.Markdown object>
```

## Problem 4: Testing properties of integers (5 points)
In this problem, you’ll get a bit more practice working with conditionals, and a ﬁrst exposure to the kind of thinking that is required in a typical “coding interview” question. A positive integer $n$ is a power of 2 if $n = 2^p$ for some integer $p \geq 0$.

### Part 1 (2 points) 
Write a function `is_power_of_2` that takes a positive integer as its only argument and returns a Boolean indicating whether or not the input is a power of 2. 

You may assume that the input is a positive integer. You may not use the built-in `math.sqrt` function in your solution. You should need only the division and modulus (%) operations. 

**Hint:** the simplest solution to this problem makes use of recursion, though recursion is not necessary.

<!-- grade_id: is_power_of_2, points: 0, grade: false, solution: true -->
```python
def is_power_of_2(n):
    
    while n % 2 == 0:
        n = n / 2

    return n == 1
```

<!-- grade_id: is_power_of_2_test1, points: 1, grade: true, solution: false -->
```python
assert is_power_of_2(1)
assert is_power_of_2(2)
assert not is_power_of_2(15)
```

<!-- grade_id: is_power_of_2_test2, points: 1, grade: true, solution: false -->
```python
assert is_power_of_2(4)
```

### Part 2 (3 points)
Generalize your previous solution to a function `is_power` that takes two positive integers as its arguments, `b` and `n`, in that order, and returns a Boolean. `is_power(b,n)` should return `True` if `n` is a power of `b` and `False` otherwise.

<!-- grade_id: is_power, points: 0, grade: false, solution: true -->
```python
def is_power(b,n):
    # YOUR CODE HERE
    if n <=0 or b <= 1:
        return False
    
    while n % b == 0:
        n = n / b

    return n == 1
```

<!-- grade_id: is_power_test1, points: 1.5, grade: true, solution: false -->
```python
assert is_power(5, 125)
assert not is_power(5, 150)
```

<!-- grade_id: is_power_test2, points: 1.5, grade: true, solution: false -->
```python
assert not is_power(1, 2)
assert is_power(3, 1)
```

## Problem 5: Create and share your stats 507 gitHub repository (5 points)

Create a new GitHub repository for the Stats 507 course and complete the following tasks:
### Set up the repository
Name the repository appropriately (e.g., "stats507-coursework");
Ensure the repository is public (unless otherwise specified)

### Add a README file
Create a file named "README.md" in the root directory;
Include a brief description of the course and the purpose of the repository

### Provide the repository link
Note: You should only post HW after the due date (or make it private)

<!-- grade_id: cell-9a484821ec48c9fc, points: 0, grade: false, solution: true -->
```python
# YOUR REPO LINK HERE
my_github_repo = 'https://github.com/zkl2002/Stats-507'
```